```python
NAME = "Shreya Paturi"
COLLABORATORS = ""
```

---

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Module 1 Assessed Exercise: Writing your first program

This notebook template includes two sections, one for each method you are going to develop. To help you get started, the first section break the potential solution into three more managable components: horizontal code calculation, vertical code calculation, and code concatenation. You should write code within each component (between 'begin solution' and 'end solution') to make them functioning. Each functional component can then be tested with the testing code following each function. Don't worry if you don't understand the testing code, it's only here to help you check if your code is functioning properly. Passing a test does not mean your code is 100% accuracte, your code will be tested more extensively once submitted. 

The second section gives you more freedom in implementing an alternative method. Please do have a try if you are aiming for more than 80%. 


## Method 1

<!-- grade_id: transx, points: 0, grade: false, solution: true -->
```python
def transformX(longitude):
    '''return the UTM horizontal code as a string given the input longitude as a float number. 
    Do not worry about special cases for the moment'''
    
    #This equation that converts the input longitude to the respective UTM zone number
    #as UTM zones are not represented in float type, the result is rounded to 0 decimal places  and also converted to integer type for a whole number, then converted to string to use it further concatenation with latitude bands
    
    utmX=str(int(round((longitude+180)/6 ,0)))
        
    return utmX# returns the converted longitude value if the function transformX is called with the longitude argument
    raise NotImplementedError()
```

<!-- grade_id: transx_tests, points: 1, grade: true, solution: false -->
```python
'''test your implementation'''
assert transformX(-163)=='3' or transformX(-163)=='03', "longitude -163 should be transformed to '3'"
print('success!')
```

Output:
```
success!
```

<!-- grade_id: transy, points: 0, grade: false, solution: true -->
```python
def transformY(latitude):
    '''return the UTM vertical code as a string given latitude as a float number. 
    Do not worry about A,B,Y or Z for the moment '''
    # YOUR CODE HERE
    #Each UTM zone is segmented into 20 latitude bands. Each latitude band is 8 degrees high with some exceptions, and is lettered starting from "C" at 80°S, increasing up to "X", omitting the letters "I" and "O" (because of their similarity to the numerals one and zero)
    #create a variable latBands to define the string containing the 20 latitude band alphabets
    #both the complete string latBands or directly returning the respective string in the if-else statements work. Either methods can be used
    latBands="CDEFGHJKLMNPQRSTUVWX"
  
    # equation((latitude+84)/8) for converting latitude to UTM latitude band does not work as there are irregularities in the difference in bands. For example, the gap between W and X band is 12 degrees, whereas, most of the bands 8 degrees apart
    #Therefore, a series of if statements are used that directly provide ranges for the input latitude values
   
    if latitude <= -72.0  and  latitude > -80.0:
        return latBands[0]        #returns string "C"
    elif latitude <= -64.0  and  latitude > -72.0:
        return latBands[1]       #returns string "D"
    elif latitude <= -56.0  and  latitude > -64.0:
        return latBands[2]      #returns string "E"
    elif latitude <= -48.0  and   latitude > -56.0:
        return latBands[3]     #returns string "F"
    elif latitude <= -40.0  and  latitude > -48.0:
        return latBands[4]     #returns string "G"
    elif latitude <= -32.0  and  latitude > -40.0:
        return latBands[5]     #returns string "H"
    elif latitude <= -24.0  and  latitude > -32.0:
        return latBands[6]    #returns string "J"
    elif latitude <= -16.0  and  latitude > -24.0:
        return latBands[7]   #returns string "K"
    elif latitude <= -8.0  and  latitude > -16.0:
        return latBands[8]   #returns string "L"
    elif latitude <= 0.0  and  latitude > -8.0:
        return latBands[9]   #returns string "M"
    elif latitude <= 8.0  and  latitude > 0.0:
        return latBands[10]  #returns string "N"
    elif latitude <= 16.0  and  latitude > 8.0:
        return latBands[11]   #returns string "P"
    elif latitude <= 24.0  and  latitude > 16.0:
        return latBands[12]   #returns string "Q"
    elif latitude <= 32.0  and  latitude > 24.0:
        return latBands[13]   #returns string "R"
    elif latitude <= 40.0  and  latitude > 32.0:
        return latBands[14]   #returns string "S"
    elif latitude <= 48.0  and  latitude > 40.0:
        return latBands[15]  #returns string "T"
    elif latitude <= 56.0  and  latitude > 48.0:
        return latBands[16]  #returns string "U"
    elif latitude <= 64.0  and  latitude > 56.0:
        return latBands[17]  #returns string "V"
    elif latitude <= 72.0  and  latitude > 64.0:
        return latBands[18]  #returns string "W"
    elif latitude <= 84.0  and  latitude > 72.0:
        return latBands[19]  #returns string "X"
    
        raise NotImplementedError()
```

<!-- grade_id: transy_tests, points: 1, grade: true, solution: false -->
```python
'''test your implementation'''
assert transformY(-5.6)=='M', "latitude 5.6 should be transformed to 'M'"
print('success!')
```

Output:
```
success!
```

<!-- grade_id: utm, points: 0, grade: false, solution: true -->
```python
def transformUTM(longitude, latitude):
    '''This function combines the functionality of transformX and transformY, returns a validate UTM zone code. 
    Note: in this function you should also write code to handle irregular UTM zones: 31V, 32V, 31X, 33X, 35X, 37X'''
    
    # Create two variables to access both the latitude and longitude UTM functions created before
    utmXfunction=transformX(longitude)# takes the longitude argument from input, executes the transformX function with the input and stores the output of the function as a variable 
    utmYfunction=transformY(latitude)# takes the latitude argument from input, executes the transformY function with the input and stores the output of the function as a variable 
    
    #A series of if and else to handle irregular UTM zones
    #On the southwest coast of Norway, grid zone 32V (9° of longitude in width) is extended further west, and grid zone 31V (3° of longitude in width) is correspondingly shrunk to cover only open water. 
    #Also, in the region around Svalbard, the four grid zones 31X (9° of longitude in width), 33X (12° of longitude in width), 35X (12° of longitude in width), and 37X (9° of longitude in width) are extended to cover what would otherwise have been covered by the seven grid zones 31X to 37X. 
    #The three grid zones 32X, 34X and 36X are not used.
    
    if utmXfunction >= '-6' and utmXfunction < '3' and utmYfunction == 'V':
        return  '31V'
    if utmXfunction >= '3' and utmXfunction > '12' and utmYfunction == 'V':
        return  '32V'
    if utmXfunction >= '0' and utmXfunction < '9' and utmYfunction == 'X':
        return  '31X'
    if utmXfunction >= '9' and utmXfunction < '21' and utmYfunction == 'X':
        return  '33X'
    if utmXfunction >= '21' and utmXfunction < '33' and utmYfunction == 'X':
        return  '35X'
    if utmXfunction >= '33' and utmXfunction < '42' and utmYfunction == 'X':
        return  '37X'
    else:
        return utmXfunction+utmYfunction
    raise NotImplementedError()
    
```

<!-- grade_id: utm_tests, points: 3, grade: true, solution: false -->
```python
'''test your implementation'''
assert transformUTM(-74.006, 40.7128)=='18T', "The UTM zone of New York should be '18T'"
print('success!')
```

Output:
```
success!
```

## Method 2

<!-- grade_id: method2, points: 0, grade: false, solution: true -->
```python
#This method uses python dictionaries to lookup latitude and longitude values and output their respective keys/UTM zone numbers and bands

# sources reffered: 
# https://stackoverflow.com/questions/23999801/python-creating-multiple-lists?
# https://www.geeksforgeeks.org/python-get-key-from-value-in-dictionary/?
#https://thispointer.com/python-dictionary-with-multiple-values-per-key/?
#https://stackoverflow.com/questions/50299172/python-range-or-numpy-arange-with-end-limit-include?
#https://en.wikipedia.org/wiki/Universal_Transverse_Mercator_coordinate_system?

#import numpy library to use it for the numpy.arange() function that generates an array of values when a start value, stop value and a step(difference between each value of the array) value is provided
import  numpy as np

#creating a function to generate UTM converted longitude
def transformX2(longitude):

#rounding off the input float values for easier use in dictionaries
    long_Float_to_Int=int(round(longitude,0))

#create an empty dictionary to populate keys and values
    longDict={}
    start= -180#longitude 1 starts at -180 
    stop= -174 #and ends at -173
    step= 1# each value has a difference of one
#create a for loop to populate keys and values
# range is given from 1 to 61 as there are 60 UTM zones(longitudes)
    for j in range(1, 61):
        #the below statement generates strings of 60 keys and also their respective values using np.arange(). Example of a key is '1', list() converts array into list 
        longDict[str(j)] = list(np.arange(start,stop,step))#LHS command for keys and RHS for values
        start=start+6 # updates the start by adding 6 as each longitude differs by a 6 degree difference
        stop=stop+6 # updates the start by adding 6 as each longitude differs by a 6 degree difference
        
    #print(longDict) #to check the keys and values of the dictionary that stores the ranges of longitudes and their corresponding UTM zone number     
    #create empty string key variable for the longitudes
    longKey=''
    # the for statement checks the keys(key) and values(list_of_values) from the items(longDict.items()) from the longDIct dictionary
    for key, list_of_values in longDict.items(): 
        #checks if the input longitude value (converted to int) exists in the values of the dictionary
            if long_Float_to_Int in list_of_values:
                #assigns the corresponding key to longKey variable
                longKey=key
    
    #the function transformX2 finally returns the UTM zone number(longKey)
    return(longKey)
    raise NotImplementedError()
    
#creating a function to generate UTM band from latitude         
def transformY2(latitude):

    #rounding off the input float values for easier use in dictionaries
    lat_Float_to_Int=int(round(latitude,0))   
    
    #create an empty dictionary to populate keys and values
    latDict={}
    #first creating a dictionary values until the key W. This is done as the number of values in X is irregualr compared to the rest of the latitude bands
    start= -80 #latitude 1 starts at -80
    stop= -72#and ends at -71
    step= 1
    #a tuple is created for storing the latitude string bands as a string variable cannot be called in a dictionary to populate the keys
    latBands=("C","D","E","F","G","H","J","K","L","M","N","P","Q","R","S","T","U","V","W","X")
    for j in range(0, 19):
        #the below statement generates strings of 19 keys and also their respective values using np.arange(). Example of a key is '1', list() converts array into list 
        latDict[latBands[j]] = list(np.arange(start,stop,step))#LHS command for keys and RHS for values
        start=start+8  # updates the start by adding 8 as each latitude differs by a 8 degree difference
        stop=stop+8  # updates the stop by adding 6 as each latitude differs by a 8 degree difference
    
    tupleX=tuple(range(72,85))# creates a tuple that stores the values for X latitude band
    latDict_1={'X':tupleX}#creates a new dictionary for X band with values of tupleX
    latDict.update(latDict_1)# this statement joins the previous dictionary latDict with latDict_1. Here latDict is updated with the new values from latDict_1
    #print(latDict) #to check the keys and values of the dictionary that stores the ranges of latitudes and their corresponding latitude band     
    
    #create empty string key variable for the latitudes
    latKey=''
     # the for statement checks the keys(key) and values(list_of_values) from the items(latDict.items()) from the latDict dictionary
    for key, list_of_values in latDict.items():
         #checks if the input latitude value (converted to int) exists in the values of the dictionary
            if lat_Float_to_Int in list_of_values:
                 #assigns the corresponding key to latKey variable
                latKey=key
    
     #the function transformY2 finally returns the UTM latitude band(latKey)
    return(latKey)
    raise NotImplementedError()

def transformUTM2(longitude, latitude):
    '''returns UTM zone code of a longitude, latitude pair. Both longitude and latitude are float numbers'''
   # Create two variables to access both the latitude and longitude UTM functions created before
    utmXfunction=transformX2(longitude)# takes the longitude argument from input, executes the transformX2 function with the input and stores the output of the function as a variable 
    utmYfunction=transformY2(latitude)# takes the latitude argument from input, executes the transformY2 function with the input and stores the output of the function as a variable 
    return utmXfunction+utmYfunction
    raise NotImplementedError()
    
```

<!-- grade_id: method2_tests, points: 2, grade: true, solution: false -->
```python
'''test your implementation'''
assert transformUTM2(144.9631,-37.8136)=='55H', "The UTM zone of Melbourne should be '55H'"
print('success!')
```

Output:
```
success!
```

```python

```

```python

```

```python

```